## Python Crash Course Exercises
Benjamin 01/06/2018

```python
from nose.tools import assert_equal
import numpy as np
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

## Exercises

Answer the questions or complete the tasks outlined in bold below, use the specific method described if applicable.

** Write a function that returns the value of a number to the power 4  **

<!-- grade_id: cell-b5e4ea39cd9bfdd2, points: 0, grade: false, solution: true -->
```python
def power_four(x):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5e8219a542924836, points: 2, grade: true, solution: false -->
```python
assert_equal(power_four(-4), 256)
```

**Write a function that splits a string and returns a list**

<!-- grade_id: cell-c749d28f1f95a54e, points: 0, grade: false, solution: true -->
```python
def split_string(x):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f00b08c30bef7cb8, points: 4, grade: true, solution: false -->
```python
assert_equal(split_string('hey there Sam'), ['hey','there','Sam'])
assert_equal(split_string('As above so below'), ['As', 'above', 'so', 'below'])
```

** Given the variables:**

    planet = "Earth"
    diameter = 12742

** Use .format() to print the following string: **

    The diameter of Earth is 12742 kilometers.

<!-- grade_id: cell-1bacb2ec9f2666af, points: 0, grade: false, solution: true -->
```python
planet = "Earth"
diameter = 12742

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-cd1c9523ccbbc77a, points: 2, grade: true, solution: false -->
```python
assert_equal(Pla_diam,'The diameter of Earth is 12742 kilometers.')
```

** Given this nested list, use indexing to grab the word "hello" **

<!-- grade_id: cell-514306edd3880e4c, points: 0, grade: false, solution: true -->
```python
lst = [1,2,[3,4],[5,[100,200,['hello']],23,11],1,7]
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-31e778fb28589d7e, points: 2, grade: true, solution: false -->
```python
assert_equal(gra,'hello')
```

** Given this nest dictionary grab the word "hello". Be prepared, this will be annoying/tricky **

<!-- grade_id: cell-27566f5cd57cdc37, points: 0, grade: false, solution: true -->
```python
d = {'k1':[1,2,3,{'tricky':['oh','man','inception',{'target':[1,2,3,'hello']}]}]}
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3bab7f23a4401347, points: 2, grade: true, solution: false -->
```python
assert_equal(nestdict,'hello')
```

**What is the main difference between a tuple and a list**

<!-- grade_id: cell-620492ccd03a78cf, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

** Create a function that returns the email website domain from a string in the form: **

    user@domain.com
    
**So for example, passing "user@domain.com" would return: domain.com**

<!-- grade_id: cell-d54c805b1093644e, points: 0, grade: false, solution: true -->
```python
def domainGet(email):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b9b6c8c34dd65caf, points: 4, grade: true, solution: false -->
```python
assert_equal(domainGet('user@domain.com'),'domain.com')
assert_equal(domainGet('kk@ymail.com'),'ymail.com')
```

** Create a basic function that returns True if the word 'dog' is contained in the input string. Don't worry about edge cases like a punctuation being attached to the word dog, but do account for capitalization. **

<!-- grade_id: cell-ffef372a0ced3059, points: 0, grade: false, solution: true -->
```python
def findDog(st):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-492d4f7da6043353, points: 5, grade: true, solution: false -->
```python
assert_equal(findDog('Is there a dog here?'), True)

```

** Create a function that counts the number of times the word "dog" occurs in a string. Again ignore edge cases. **

<!-- grade_id: cell-c4dca3108ecbb2c6, points: 0, grade: false, solution: true -->
```python
def countDog(st):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e1b751fa8d5e84a9, points: 2, grade: true, solution: false -->
```python
assert_equal(countDog('This dog runs faster than the other dog dude!'),2)
```

** Use lambda expressions and the filter() function to filter out words from a list that don't start with the letter 's'. For example:**

    seq = ['soup','dog','salad','cat','great']

**should be filtered down to:**

    ['soup','salad']

<!-- grade_id: cell-14c8834296ad5773, points: 0, grade: false, solution: true -->
```python
def filter_list(seq):
     # YOUR CODE HERE
     raise NotImplementedError()
```

<!-- grade_id: cell-516be67e311c669e, points: 5, grade: true, solution: false -->
```python
assert_equal(filter_list(['soup','dog','salad','cat','great']), ['soup','salad'])
```

```python

```